---
title       : Modulo 2. Introduccion alLenguaje R
description : Este es un capitulo introductorio para dejar claras las bases del lenguaje R.

--- type:NormalExercise lang:r xp:100 skills:1 key:88e5a6a670
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
