---
title       : Modulo 3. xxxxxx
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:d6785c802f
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
